```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import numpy as np
import numpy.testing as nt
```

# Optimisation II - Test 2


## Instructions

* You may only access https://moodle.ms.wits.ac.za/tests/ during the test and **NOT** normal moodle (i.e. https: //moodle.ms.wits.ac.za/moodle/ )
* All machines are being logged and anyone caught logging into the regular Moodle listed above will be given **zero**. 
* Anyone caught trying to log into someone else's account will be given **zero**.
* Anyone attempting to use brought in code, notes or any other materials will be given **zero**.
* Read all the instructions carefully.
* **DO NOT change the name of the notebook in any way!**
* The test will run from 08:15 - 10:00 (this includes the extra time students). At 10:00 Moodle will no longer take submissions so ensure that you have uploaded your notebook by at least 11:55.
* DO NOT rename the file. Simply answer all questions within the notebook and resubmit the file.
* There are 105 marks available. There are 5 MCQ questions and 5 programming questions.

## MCQ Questions

If each of the MCQ questions there is only one correct answer. Once you have decided a answer in each question, uncomment out the response of your choosing from the correspond python function.

There are 5 MCQ questions.

***

### Question 1

Applying Newton's method on the quadratic function: $$f(\mathbf{x}) = \dfrac{1}{2}\mathbf{x}^TQ\mathbf{x} + \mathbf{x}^T \mathbf{b} + \mathbf{c},$$where $Q$ is a symmetric positive-definite matrix, will converge to the minimum in how many steps?

* (a) 1
* (b) Depends on the starting point
* (c) Depends on the tolerance used
* (d) Depends on the loops determined by the *for-loop*
* (e) None of the above

<!-- grade_id: cell-f0ffb7bf7938c1ba, points: 0, grade: false, solution: true -->
```python
def question1():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    #return 'e'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5190bc15807307bd, points: 0, grade: true, solution: false -->
```python
assert(question1() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-73bc8f21efa81efd, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Hidden Evaluation of Answer
```

### Question 2

The steepest descent method makes use of what when considering a search direction:

* (a) No derivatives are used
* (b) The first derivative is used
* (c) Both the first and second derivatives are used
* (d) The second derivative is used
* (e) None of the above

<!-- grade_id: cell-fad53ad658a04108, points: 0, grade: false, solution: true -->
```python
def question2():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    #return 'e'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-51e8d01b2a97ab11, points: 0, grade: true, solution: false -->
```python
assert(question2() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-c8e39f37cca5beba, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Hidden Evaluation of Answer
```

### Question 3

Given the Rosenbrock function:
$$
f(x)=100(x_2-x_1^2)^2+(1-x_1)^2.
$$
Applying Newton's method with a starting value of $\mathbf{x}^0= {0\choose 0}$ yields the next iterate $\mathbf{x}^1= {1\choose 0}$. What can you conclude about the performance of the algorithm:

* (a) The algorithm is converging to a minimum
* (b) Nothing, more iterations need to be performed in order for a patten to emerge
* (c) Nothing, the patten of iterates is random
* (d) The algorithm is diverging
* (e) None of the above

<!-- grade_id: cell-ec17e576119d50a1, points: 0, grade: false, solution: true -->
```python
def question3():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    #return 'e'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d9f98220c2b345ff, points: 0, grade: true, solution: false -->
```python
assert(question3() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-68f89d3eb95965db, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Hidden Evaluation of Answer
```

### Question 4

Consider the Modified Newton Method. One major drawback of the method is:

* (a) The method has poor convergence and is slow
* (b) The method rarely works on non-quadratic functions
* (c) The method is computational expensive due to computing the inverse Hessian per iterate
* (d) The method can diverge even if $H^{k}$ is positive definite
* (e) None of the above

<!-- grade_id: cell-5276b192c0be54da, points: 0, grade: false, solution: true -->
```python
def question4():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    #return 'e'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-316ef23e7fc216f7, points: 0, grade: true, solution: false -->
```python
assert(question4() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-dd6f73c25a7dc1e4, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Hidden Evaluation of Answer
```

### Question 5

Random Walk and Downhill Simplex use which derivative(s) when computing a new iterate:

* (a) First
* (b) Second 
* (c) Second but First used in a line search
* (d) None
* (e) None of the above

<!-- grade_id: cell-1561778c493f358f, points: 0, grade: false, solution: true -->
```python
def question5():
    """
    Uncomment your choice to the question
    """
    #return 'a'
    #return 'b'
    #return 'c'
    #return 'd'
    #return 'e'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a359c2e1235d8f2c, points: 0, grade: true, solution: false -->
```python
assert(question5() in ['a', 'b', 'c', 'd', 'e'])
print('Function is returning a plausible response')
```

<!-- grade_id: cell-f2a7af867c36da82, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Hidden Evaluation of Answer
```

## Programming Questions 

### Question  1 [15 Marks]

Converting the polar coordinate $(r, \theta)$ of a point to the Cartesian coordinates $(x, y)$ we have:
\begin{eqnarray*}
x = r\cos\theta,\\
y = r\sin\theta,\\
r = x^2 + y^2.
\end{eqnarray*}
Write a function `neighbourhood`, which takes as input a origin point `c` (i.e. ($x_0, y_0$), as a numpy array), `r` a radial distance and `n` scalar number of samples. The function should generate `n` samples in the neighbourhood around `c`. The samples should be generated randomly within the radius `r`.

<!-- grade_id: cell-a2b5fa6bc5b4d111, points: 0, grade: false, solution: true -->
```python
def neighbourhood(c, r, n):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-e763ce7265cca916, points: 15, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 15 Marks
import numpy.linalg as LA
c = [np.random.randint(10), np.random.randint(10)]
r = np.random.randint(10)
n = 2000
xtest, ytest = neighbourhood(c, r, n)
for i in range(n):
    pt   = np.array([xtest[i], ytest[i]])
    dist = LA.norm(pt - c, 2)
    if dist > r:
        raise AssertionError('point exceeds maximum Euclidean distance')
print('Test case passed!!!')


```

### Question  2 [15 Marks]

Write a function `secant_minimiser`, which takes as inputs, the function `f`, its derivative `g`, initial guesses, `x0` and `x1`. It should also take in the scalar `tol`. The function should perform the secant method to determine the minima of `f` as well as the corresponding $x$ value associated with the minima. You must use `abs(x1 - x0)` to control your tolerance.

**Recall:** The updating formula for the *standard Secant Method for root finding* is:
$$
x_{n+1}=x_n-f(x_n)\dfrac{(x_n-x_{n-1})}{\left (f(x_n)-f(x_{n-1})
\right)}. \nonumber
$$

<!-- grade_id: cell-584640c8cd53a6b4, points: 0, grade: false, solution: true -->
```python
def secant_min(f, g, x0, x1, tol):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4c376330935755c5, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 2 Marks
f   = lambda x: x**2 - 612
g   = lambda x: 2*x
x0  = 10
x1  = 30
tol = 1e-5
secant_min(f, g, x0, x1, tol)
nt.assert_almost_equal(-612.0, secant_min(f, g, x0, x1, tol)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-554cff2e99a8a255, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 2 Marks
f   = lambda x: x**2 - 612
g   = lambda x: 2*x
x0  = 10
x1  = 30
tol = 1e-5
secant_min(f, g, x0, x1, tol)
nt.assert_almost_equal(0, secant_min(f, g, x0, x1, tol)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-b59deab4e43649bc, points: 5, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 5 Marks
```

<!-- grade_id: cell-b31a54215541f045, points: 6, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 6 Marks
```

### Question  3 [20 Marks]

Write a function `steepest_descent` which performs two variants of the method of steepest descent. The function should take in as inputs the function, `f`, its derivative `g`, an initial guess `x0`, a tolerance `tol`, a backtracking parameter `beta` and optional input `alpha`. If the optional input `alpha` is passed, then the function should perform the standard method of steepest descent with the given fixed step-size `alpha`. If no optional input is passed, then the function must perform the steepest descent with backtracking line search. You must use `beta` to find the acceptable `alpha` at each iterate. Both variants of the function should return the approximate minimum $x^*$, as well as the number of iterations required to achieve it. The logic to control your main while loops should be the `norm(g(x0), 2)`. In the case of backtracking, recall the while loop control:
$$
f(x - t\nabla f(x)) > f(x) - \dfrac{t}{2}\lVert \nabla f(x) \rVert^2.\nonumber
$$

**Note:** it would be advisable to limit your while limits in case of an infinite loop.

<!-- grade_id: cell-d7f7eba75cff492e, points: 0, grade: false, solution: true -->
```python
def steepest_descent(f, g, x0, tol, beta, *alpha):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a85b42d6c4291ddb, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 1 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-7
alpha = 8/70
nt.assert_array_almost_equal([2.02246559e-08, 2.64493244e-15], steepest_descent(f1, g1, x0, tol, beta, alpha)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-3947dd80c23ad8e2, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 1 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-7
alpha = 8/70
assert(abs(steepest_descent(f1, g1, x0, tol, beta, alpha)[1] - 29) < 2)
print('Test case passed!!!')
```

<!-- grade_id: cell-aec54515523b6b0c, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 1 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-7
alpha = 8/70
nt.assert_array_almost_equal([ 3.42518076e-09, -1.39964061e-10], steepest_descent(f1, g1, x0, tol, beta)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-570d58288c1b89cc, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 1 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-7
alpha = 8/70
assert(abs(steepest_descent(f1, g1, x0, tol, beta)[1] - 11) < 2)
print('Test case passed!!!')
```

<!-- grade_id: cell-07ceaf2128b6a435, points: 4, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 4 Marks
```

<!-- grade_id: cell-955e1736500943ad, points: 4, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 4 Marks
```

<!-- grade_id: cell-42ef9d86fc919c88, points: 4, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 4 Marks
```

<!-- grade_id: cell-2b5601dae5d9ab75, points: 4, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 4 Marks
```

### Question  4 [15 Marks]

Recall the algorithm for the DFP Method given below:

1. Set $k = 0, G^0 = I$ and compute $\mathbf{g}^k = g(\mathbf{x}^k)$.
2. Compute $\mathbf{d}^k$ from $\mathbf{d}^k = -G^k\mathbf{g}^k$.
3. Compute $\alpha = \alpha^k$ such that $f(\mathbf{x}^k + \mathbf{\alpha}^k\mathbf{d}^k),$ set $\mathbf{x}^{k+1} = \mathbf{x}^{k} + \alpha^k\mathbf{d}^k$.
4. Compute $\mathbf{g}^{k+1}$ such that $\mathbf{g}^{k+1} = g(\mathbf{x}^{k+1})$.
5. If $\lVert g^{k+1} \rVert \leq \epsilon$ ($\epsilon$ is a user supplied small number) then go to (9).
6. Compute $\delta^k$ and $\gamma^k$ such that $\delta^k = \mathbf{x}^{k+1} - \mathbf{x}^k$ and $\gamma = \mathbf{g}^{k+1} - \mathbf{g}^k$.
7. Compute $G^{k+1}$.
8. Set $k = k + 1$ and go to (2).
9. Set $\mathbf{x}^* = \mathbf{k+1}$, STOP.

The updating rank-2 formula is given below:
\begin{equation*}
G^{k+1}=G^k+\dfrac{\delta^k \delta^{k^T}}{\delta^{k^T}\gamma^k} - \dfrac{G^k\gamma^k{\gamma^k}^TG^k}{{\gamma^k}^T G^k\gamma^k}   
\end{equation*}

Finish the implementation seen below according to the above. The function should return the final approximation, the number of iterations and the final $G^k$.

<!-- grade_id: cell-0806f3a1c2f5ad41, points: 0, grade: false, solution: true -->
```python
def dfp(f, g, x0, tol, beta):
    import numpy.linalg as LA
    G  = np.eye(len(x0))
    it = 0
    while LA.norm(g(x0)) > tol and it < 20:
        d = -G.dot(g(x0))
        alpha = 1.
        while f(x0 + np.dot(alpha,d)) > f(x0) - (alpha/2)*(LA.norm(d)**2):
            alpha = beta*alpha        
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e86cfe738a38f123, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 2 Marks
f2    = lambda x: x[0] - x[1] + 2*x[0]**2 + 2*x[0]*x[1] + x[1]**2
g2    = lambda x: np.array([1 + 4*x[0] + 2*x[1], -1 + 2*x[0] + 2*x[1]])
x0    = np.array([0, 0])
beta  = 0.9
tol   = 1e-7

t1 = np.array([-1. ,  1.5])
t2 = 9
t3 = np.array([[ 0.50165299, -0.50078642],
               [-0.50078642,  1.00037415]])
nt.assert_array_almost_equal(t1, dfp(f2, g2, x0, tol, beta)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-3d6a4b2a385d7dd7, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 2 Marks
nt.assert_almost_equal(t2, dfp(f2, g2, x0, tol, beta)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-16352b9e48c7d127, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
# 2 Marks
nt.assert_array_almost_equal(t3, dfp(f2, g2, x0, tol, beta)[2])
print('Test case passed!!!')
```

<!-- grade_id: cell-e12ad5d8b017ea91, points: 3, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 3 Marks
```

<!-- grade_id: cell-ff4164a0ac391296, points: 3, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 3 Marks
```

<!-- grade_id: cell-7287ae597fd7d60d, points: 3, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 3 Marks
```

### Question  5 [15 Marks]

Write a function that called performs a multiplication table. The function takes in inputs, $n$ and $m$. The function outputs a $n\times m$  matrix which is the resultant multiplication table. That is, the element at row $i$ and column $j$ equals to $i\times j$. The function is also to have a second output which is the sum of all elements of the matrix.

<!-- grade_id: cell-9b713c7c9b0ae5c0, points: 0, grade: false, solution: true -->
```python
def multi_table(n, m):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b2db29d3b9799b48, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Mark
o1 = np.array([[  1.,   2.,   3.,   4.,   5.,   6.,   7.,   8.,   9.,  10.],
               [  2.,   4.,   6.,   8.,  10.,  12.,  14.,  16.,  18.,  20.],
               [  3.,   6.,   9.,  12.,  15.,  18.,  21.,  24.,  27.,  30.],
               [  4.,   8.,  12.,  16.,  20.,  24.,  28.,  32.,  36.,  40.],
               [  5.,  10.,  15.,  20.,  25.,  30.,  35.,  40.,  45.,  50.],
               [  6.,  12.,  18.,  24.,  30.,  36.,  42.,  48.,  54.,  60.],
               [  7.,  14.,  21.,  28.,  35.,  42.,  49.,  56.,  63.,  70.],
               [  8.,  16.,  24.,  32.,  40.,  48.,  56.,  64.,  72.,  80.],
               [  9.,  18.,  27.,  36.,  45.,  54.,  63.,  72.,  81.,  90.],
               [ 10.,  20.,  30.,  40.,  50.,  60.,  70.,  80.,  90., 100.],
               [ 11.,  22.,  33.,  44.,  55.,  66.,  77.,  88.,  99., 110.],
               [ 12.,  24.,  36.,  48.,  60.,  72.,  84.,  96., 108., 120.]])

nt.assert_array_equal(o1, multi_table(12, 10)[0])
print('All Tests Passed!!!')
```

<!-- grade_id: cell-08c0ad32aef0fe62, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 12 Mark
nt.assert_array_equal(4290, multi_table(12, 10)[1])
print('All Tests Passed!!!')
```

<!-- grade_id: cell-e6cacbe6fcfb8f67, points: 6, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 6 Marks
```

<!-- grade_id: cell-8b529db97ede3e03, points: 5, grade: true, solution: false -->
```python
# Do not delete this cell
# Hidden test
# 5 Marks
```