```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
```

# Optimisation II - Lab 5


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!
* To launch jupyter correctly in the labs using vmuser, type the following at the terminal:
    `/usr/local/anaconda3/bin/jupyter-notebook`

## Main Exercises

Complete the following problems:



### Question 1

Write a function `downhill_simplex` that takes as inputs, a multivariate function `f`, an initial guess `x0`, a maximum number of iterations `it` along with values for `alpha`, `gamma` and `beta`. The function should return the minimum of the function `f`.

<!-- grade_id: cell-b216aec15abddf54, points: 0, grade: false, solution: true -->
```python
def downhill_simplex(f, x0, it, alpha = 1, gamma = 2, beta = 0.5):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-bcbf7f713f18c4cb, points: 10, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 10 Marks
x0    = np.array([[-1.20, 1.00], [-0.23, 1.26], [ -0.94, 1.97]])
f     = lambda x: (1 - x[0])**2 + 1000*(x[1] - x[0]**2)**2
alpha = 1
beta  = 0.5
gamma = 2
it    = 200
downhill_simplex(f, x0, it, alpha, gamma, beta)
t1 = np.array([[1., 1.],
               [1., 1.],
               [1., 1.]])
nt.assert_almost_equal(t1, downhill_simplex(f, x0, it, alpha, gamma, beta))
print('Test case passed!!!')
```

<!-- grade_id: cell-ad2ef526a5f06fd2, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
# Hidden cases
```


### Question 2

Write a function `random_walk` which takes as inputs, a multivariate function `f`, an initial starting point `x`, a step length `lambda`, a tolerance for $\lambda$ `eps`, the dimension of $f$ `n` and a maximum number of iterations `it`. The function should return the approximation to the minima as well as the function value at this point, i.e. `x` and `f(x)`.

<!-- grade_id: cell-d4632b35e783f761, points: 0, grade: false, solution: true -->
```python
def random_walk(f, x, l, eps, n, it):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5509b3f10f10cbfe, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f   = lambda x1, x2: x1 - x2 + 2*x1**2 + 2*x1*x2 + x2**2
x0  = np.array([0, 0])
lam = 1
eps = 1e-5
n   = 2 
N   = 200
nt.assert_almost_equal([-1.00000726,  1.50002065], random_walk(f, x0, lam, eps, n, N)[0], decimal=4)
print('Test case passed!!!')
```

<!-- grade_id: cell-193125148c5a4175, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f   = lambda x1, x2: x1 - x2 + 2*x1**2 + 2*x1*x2 + x2**2
x0  = np.array([0, 0])
lam = 1
eps = 1e-5
n   = 2 
N   = 200
nt.assert_almost_equal(-1.2499999999999898, random_walk(f, x0, lam, eps, n, N)[1])
print('Test case passed!!!')
```