```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
```

# Optimisation II - Lab 4


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!
* To launch jupyter correctly in the labs using vmuser, type the following at the terminal:
    `/usr/local/anaconda3/bin/jupyter-notebook`

## Main Exercises

Complete the following problems:

### Question 1

Write a function `steepest_descent` which implements the method of Steepest Descent using a backtracking line search. The function should take as inputs the multivariate function `f`, the gradient `g`, some initial guess `x`, some dampening factor `beta` and a tolerance `tol`. Use `norm(x, 2)` to control your tolerance. Use the same control as the notes for the backtracking, i.e: 
$$f(x - t\nabla f(x)) > f(x) - \dfrac{t}{2}\lVert \nabla f(x) \rVert^2,$$
The function should return three outputs, the final approximation `x`, the function value at said `x`, i.e. `f(x)` and the number of iterations required.

<!-- grade_id: cell-e34f89d9c943a4e0, points: 0, grade: false, solution: true -->
```python
def steepest_descent(f, g, x0, beta, tol):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ec1a6bb300de6eb8, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-5
nt.assert_array_almost_equal(4.510141461145487e-13, steepest_descent(f1, g1, x0, beta, tol)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-588e8ff948be164a, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-5
nt.assert_array_almost_equal([ 4.74874245e-07, -1.00812235e-09], steepest_descent(f1, g1, x0, beta, tol)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-8c956c75e3fdca4a, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([4*x[0], 6*x[1]])
x0    = np.array([1, 1])
beta  = 0.9
tol   = 1e-5
nt.assert_array_almost_equal(9, steepest_descent(f1, g1, x0, beta, tol)[2])
print('Test case passed!!!')
```


### Question 2

Write a function that performs the DFP algorithm. The function should take as input, the function `f`, the derivative `g`, an initial starting point `x0`, a tolerance `tol` and a fixed step-size `alpha`. The function should return the final approximation and the number of iterations required.

<!-- grade_id: cell-ddb4e54114c0b525, points: 0, grade: false, solution: true -->
```python
def dfp(f, g, x0, tol, alpha):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-afc73c3e249d16fe, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f2    = lambda x: x[0] - x[1] + 2*x[0]**2 + 2*x[0]*x[1] + x[1]**2
g2    = lambda x: np.array([1 + 4*x[0] + 2*x[1], -1 + 2*x[0] + 2*x[1]])
x0    = np.array([0, 0])
beta  = 0.75
tol   = 1e-7
nt.assert_almost_equal([-1.        ,  1.50000001], dfp(f2, g2, x0, tol, beta)[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-320386cda33a630d, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f2    = lambda x: x[0] - x[1] + 2*x[0]**2 + 2*x[0]*x[1] + x[1]**2
g2    = lambda x: np.array([1 + 4*x[0] + 2*x[1], -1 + 2*x[0] + 2*x[1]])
x0    = np.array([0, 0])
beta  = 0.75
tol   = 1e-7
assert(14 == dfp(f2, g2, x0, tol, beta)[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-31739549b1c09af9, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
# Hidden cases
```

<!-- grade_id: cell-a49fcc6eba0f925c, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
# Hidden cases
```

```python

```